```python
import pandas as pd
```

```python
unames = ['user_id', 'gender', 'age', 'occupation', 'zip'] 
users = pd.read_table('users.dat', sep='::', header=None, names=unames, engine='python', index_col='user_id')
```

```python
mnames = ['movie_id', 'title', 'genres']
movies = pd.read_table('movies.dat', sep='::', header=None, names=mnames, engine='python', index_col='movie_id')
```

```python
ratings = pd.read_table('ratings.csv', sep=',', index_col='id')
```

---

# Groupby

Q. ratings를 movie 별로 묶어서 각 영화마다 ratings의 평균을 구해보세요. 

<!-- grade_id: A0, points: 0, grade: false, solution: true -->
```python
# A0 = 
### BEGIN SOLUTION
A0 = ratings.groupby('movie')['rating'].mean()
### END SOLUTION
```

```
movie
1    4.333333
2    2.333333
3    4.000000
5    4.000000
6    3.500000
Name: rating, dtype: float64
```

<!-- grade_id: A0_test, points: 1, grade: true, solution: false -->
```python
assert len(A0) == 1770
assert A0.name == 'rating'
assert round(A0.mean(), 2) == 3.42
```

Q. ratings를 user별로 묶어서 각 유저 별 평점 평균을 구해보세요.

<!-- grade_id: A1, points: 0, grade: false, solution: true -->
```python
# A1 = 
### BEGIN SOLUTION
A1 = ratings.groupby('user')['rating'].mean()
### END SOLUTION
```

```
user
2783    3.0
2785    3.0
2786    4.0
2788    4.0
2790    4.0
Name: rating, dtype: float64
```

<!-- grade_id: A1_test, points: 1, grade: true, solution: false -->
```python
assert len(A1) == 1987
assert A1.index.name == 'user'
assert round(A1.mean(), 2) == 3.66
```

Q. ratings에서 각 rating 점수 별로 몇 번 평가가 이루어졌는지 구해보세요.

<!-- grade_id: A2, points: 0, grade: false, solution: true -->
```python
# A2 = 
### BEGIN SOLUTION
A2 = ratings.groupby('rating').size()
### END SOLUTION
```

```
rating
1     305
2     487
3    1321
4    1757
5    1130
dtype: int64
```

<!-- grade_id: A2_test, points: 1, grade: true, solution: false -->
```python
assert len(A2) == 5
assert A2.sum() == 5000
assert A2.index.name == 'rating'
```

Q. users를 gender 별로 묶어서 여자와 남자가 각각 몇 명이 있는지 구해보세요.

<!-- grade_id: A3, points: 0, grade: false, solution: true -->
```python
# A3 = 
### BEGIN SOLUTION
A3 = users.groupby('gender').size()
### END SOLUTION
```

```
gender
F    1709
M    4331
dtype: int64
```

<!-- grade_id: A3_test, points: 1, grade: true, solution: false -->
```python
assert A3.index.name == 'gender'
assert 'F' in A3.index and 'M' in A3.index
assert A3.sum() == 6040
```

Q. users에서 gender 별로 age의 평균이 얼마인지 구해보세요.

<!-- grade_id: A4, points: 0, grade: false, solution: true -->
```python
# A4 = 
### BEGIN SOLUTION
A4 = users.groupby('gender')['age'].mean()
### END SOLUTION
```

```
gender
F    30.859567
M    30.552297
Name: age, dtype: float64
```

<!-- grade_id: A4_test, points: 1, grade: true, solution: false -->
```python
assert A4.index.name == 'gender'
assert 'F' in A4.index and 'M' in A4.index
assert round(A4.F, 2) == 30.86
assert round(A4.M, 2) == 30.55
```