```python
import pandas as pd
```

```python
unames = ['user_id', 'gender', 'age', 'occupation', 'zip'] 
users = pd.read_table('users.dat', sep='::', header=None, names=unames, engine='python', index_col='user_id')
```

```python
mnames = ['movie_id', 'title', 'genres']
movies = pd.read_table('movies.dat', sep='::', header=None, names=mnames, engine='python', index_col='movie_id')
```

```python
ratings = pd.read_table('ratings.csv', sep=',', index_col='id')
```

---

Q. users에서 20대 중반 이상의 유저들만 보려면? (즉, 24세 이하를 제외한다면)

<!-- grade_id: A0, points: 0, grade: false, solution: true -->
```python
# A0 =
### BEGIN SOLUTION
A0 = users[lambda x: x.age>24]
### END SOLUTION
```

```
        gender  age  occupation    zip
user_id                               
2            M   56          16  70072
3            M   25          15  55117
4            M   45           7  02460
5            M   25          20  55455
6            F   50           9  55117
```

<!-- grade_id: A0_test, points: 1, grade: true, solution: false -->
```python
assert len(A0) == 4715
assert (A0.age > 24).all()
```

Q. users에서 20대 중반 이상의 여자들만 보고 싶다면?

<!-- grade_id: A1, points: 0, grade: false, solution: true -->
```python
# A1 = 
### BEGIN SOLUTION
A1 = users[lambda x: x.age>24][lambda x: x.gender == 'F']
### END SOLUTION
```

```
        gender  age  occupation    zip
user_id                               
6            F   50           9  55117
10           F   35           1  95370
11           F   25           1  04093
16           F   35           0  20670
24           F   25           7  10023
```

<!-- grade_id: A1_test, points: 1, grade: true, solution: false -->
```python
assert len(A1) == 1333
assert (A1.gender == 'F').all()
assert (A1.age > 24).all()
```

Q. ratings에서 1번 영화를 본 유저들과 그들이 준 rating 점수를 보고 싶다면?

<!-- grade_id: A2, points: 0, grade: false, solution: true -->
```python
# A2 = 
### BEGIN SOLUTION
A2 = ratings.loc[lambda x: x.movie == 1, ['user', 'rating']]
### END SOLUTION
```

```
        user  rating
id                  
2787_1  2787       5
2788_1  2788       5
2792_1  2792       3
2796_1  2796       3
2799_1  2799       5
```

<!-- grade_id: A2_test, points: 1, grade: true, solution: false -->
```python
assert len(A2) == 15
assert 'user' in A2.columns and 'rating' in A2.columns and len(A2.columns) == 2
assert all(A2.index.str[-1] == '1')
```

Q. movies에서 movie_id가 1인 영화만 하나만 보려면?

<!-- grade_id: A3, points: 0, grade: false, solution: true -->
```python
# A3 = 
### BEGIN SOLUTION
A3 = movies.loc[1]
### END SOLUTION
```

```
title                Toy Story (1995)
genres    Animation|Children's|Comedy
Name: 1, dtype: object
```

<!-- grade_id: A3_test, points: 1, grade: true, solution: false -->
```python
assert A3.name == 1
assert A3.genres == "Animation|Children's|Comedy"
assert len(A3.index) == 2
```

Q. movies에서 movie_id가 1인 영화의 title을 보려면?

<!-- grade_id: A4, points: 0, grade: false, solution: true -->
```python
# A4 = 
### BEGIN SOLUTION
A4 = movies.loc[1, 'title']
### END SOLUTION
```

<!-- grade_id: A4_test, points: 1, grade: true, solution: false -->
```python
assert all(movies[lambda x: x.title == A4].index == 1)
```

Q. ratings에서 1점을 받은 movie만 뽑아 보려면?

<!-- grade_id: A5, points: 0, grade: false, solution: true -->
```python
# A5 = 
### BEGIN SOLUTION
A5 = ratings[lambda x: x.rating == 1]
### END SOLUTION
```

```
           user  movie  rating
id                            
2783_208   2783    208       1
2784_3248  2784   3248       1
2784_3566  2784   3566       1
2785_2396  2785   2396       1
2786_637   2786    637       1
```

<!-- grade_id: A5_test, points: 1, grade: true, solution: false -->
```python
assert len(A5) == 305
assert all(A5.rating == 1)
```

Q. ratings에서 3934번 유저의 평점들만 보려면?

<!-- grade_id: A6, points: 0, grade: false, solution: true -->
```python
# A6 = 
### BEGIN SOLUTION
A6 = ratings.loc[lambda x: x.user == 3934, 'rating']
### END SOLUTION
```

```
id
2783_1253    5
2783_589     5
2783_1270    4
2783_1274    4
2783_741     5
Name: rating, dtype: int64
```

<!-- grade_id: A6_test, points: 1, grade: true, solution: false -->
```python
assert len(A6) == 10
assert A6.mean() == 2.1
```

Q. movies에서 Toy Story (1995), Heat (1995), Copycat (1995)의 장르만 보려면?

<!-- grade_id: A7, points: 0, grade: false, solution: true -->
```python
# A7 =
### BEGIN SOLUTION
A7 = movies.loc[[1,6, 22], 'genres']
### END SOLUTION
```

```
movie_id
1     Animation|Children's|Comedy
6           Action|Crime|Thriller
22           Crime|Drama|Thriller
Name: genres, dtype: object
```

<!-- grade_id: A7_test, points: 1, grade: true, solution: false -->
```python
assert len(A7) == 3
assert A7.name == 'genres'
assert A7.iloc[0] == "Animation|Children's|Comedy"
assert A7.str.contains('Thriller').sum() == 2
```